<center>
    <h1>Image Processing I</h1>
    <h5>Fall 2020</h5>
    <h2>Introduction Laboratory</h2>
    <h4>September 17 - 1% of overall grade</h4>
</center>

This notebook is part of a project funded by the EPFL Center for Digital Education (CEDE) and the Faculty of Engineering (STI). It is owned by the Biomedical Imaging Group (BIG). Its distribution or reproduction is strictly prohibited.

### Student Name: 
### SCIPER: 

Double-click on this cell, fill your name and SCIPER number above, and run the cell below to verify your identity in Noto and set the seed for random results.

```sos
%use sos
import getpass
# This line recovers your camipro number to mark the images with your ID
uid = int(getpass.getuser().split('-')[2]) if len(getpass.getuser().split('-')) > 2 else ord(getpass.getuser()[0])
print(f'SCIPER: {uid}')
```

# <a name="top"></a> Introduction Lab (**8 Points**)

This introductory lab will guide you through the tools that you will be using in the labs for the courses Image Processing I (**IP1**) ([MICRO-511](https://moodle.epfl.ch/course/view.php?id=522)) and Image Processing II (**IP2**) (MICRO-512). As you will notice, you will be using 2 programming languages: JavaScript and Python. Don't worry if you do not have experience with them! In fact, you don't need any programming experience at all. Basic computational thinking (e.g. how *for* loops work, what a *function* is) is enough to go through the labs. We will always teach by example, and you can always go back and check an example we gave you to know how to code anything we ask. Moreover, you can always ask specific questions about programming to the dedicated TAs of the course.

<div class="alert alert-danger">

**Warning:** This lab is **long**. It is intended as a tutorial course for you to understand all the details of what happens in all the labs to come, even in the code you do not fill yourself. This implies that 
<ul>
    <li> You should not worry if it feels hard, long, or a bit too much at first,
    <li> You should keep your filled copy for further reference. 
</ul>
</div>

At the beginning of each lab, you will always find an Index as the one below, clearly pointing at the topics that will be covered. The index will also specify how many points are awarded for each exercise, when appropriate.

## Index
1. [Intro to Jupyter Notebooks](#1_intro_python) 
    1. [Cellular Structure](#1A)
    2. [The SoS Kernel](#1B)
2. [Python and JavaScript](#2_python_and_js)
    1. [Basic Python Coding](#2A)
        1. [NumPy](#2Aa)
        2. [Matplotlib](#2Ab)
    2. [Basic Javascript Coding](#2B)
    3. [Python vs JavaScript - speed comparison](#2C)
3. [IPLabImageAccess](#2_ImageAccess) (**2 Points**)
4. [IPLabViewer](#4_viewer)
    1. [Creation of a viewer](#instantiation_viewer)
    2. [Using Widgets](#widgets_viewer)
    3. [User Defined Widgets](#usr_widgets_viewer)
    4. [Programmatic Customization](#prog_customization)
    5. [Try it Yourself!](#ex_viewer) (**4 Points**)
5. [IP in Python](#5_ippython)
    1. [Examples](#examples_iplibraries)
    1. [Try it Yourself!](#exercise_iplibraries) (**2 Points**)

# <a name="1_intro_python"></a> 1. Introduction to Jupyter Notebooks
## <a name="1A"></a> 1.A. Cellular Structure
What you are looking at right now is a Jupyter Notebook. A notebook consists of cells that contain either code snippets and their results or text. This makes it the perfect place where to teach and learn image processing, as you don't need to switch between different files to see the explanation, the corresponding code, and its results. The cellular structure of Jupyter Notebook allows us to run small snippets of code one after the other, go back and forth between, or check intermediate results between them. 
<div class="alert alert-info">

**Note:** You can use the shortcut `Shift` $+$ `Enter` to run the cell and advance to the next one or `Ctrl` $+$ `Enter` to run the cell and stay in the current cell.
</div>

It is important to know that if you declare some variable in a cell, this variable will be available in all other cells that use the same programming language, irrespective of the location of the cell. Let's look at an example: Read the description of the cells below and run them as instructed.

In this first cell we declare a variable `a` and assign it the value $5$. Run it.

```sos
%use sos
# cell 1
# this is our variable
a = 5
```

In the second cell we print the value of `a` to the console. If you run it just after the cell above, it should print "The value of a is 5". Run it.

```sos
%use sos
# cell 2
# here we print the value of a
print(f"The value of a is {a}")
```

Then, in the third cell, we modify `a` by adding $2$ to it. Run it.

```sos
%use sos
# cell 3
# add two to a
a = a + 2
```

Now run _cell 2_ again: the value of `a` has been modified, even though _cell 2_ is located above _cell 3_ in our notebook. Repeat this procedure (running _cell 3_ then running _cell 2_ ) and make sure you can explain the result: every time you do `a` will be incremented by 2. Now run _cell 1_ again, where we set `a=5`, and `a` will be $5$ in all cells again. Hopefully this shows the **global** impact every single cell has on the notebook. Regardless, it is good practice to design notebooks to be run from top to bottom.

## <a NAME="1B"></a> 1.B. The SoS Kernel
Surely you have noticed that the first line of all cells with code was `%use sos`. This line tells the notebook that this specific cell uses Python syntax and should be read as Python code. In normal Jupyter Notebooks, this would not be necessary, since they usually only support one single programming language in the whole notebook, defined by _the kernel_. However, for **IP1** and **IP2**, we use a special kernel that allows us to have a notebook with multiple programming languages, the [SoS](https://vatlab.github.io/sos-docs/) kernel. If you're interested you can read its fascinating [documentation](https://vatlab.github.io/sos-docs/notebook.html#Documentation). However, to understand the labs you only need to know that:
- `%use sos` indicates that a cell is written in Python,
- `%use javascript` indicates that a cell is written in JavaScript,
- `%put var_name ` converts the variable `var_name` from JavaScript to Python
- `%get var_name ` converts the variable `var_name` from Python to JavaScript

Those keywords are called _magics_ and they always need to be on the first lines of code with nothing in-between them. Let's look at an example: Read the description of the cells below and run them as instructed.

First, we declare some variable `my_var` in python.

```sos
%use sos
# This is our variable, let's initialize it to 0
my_var = 0
```

Next, we want to modify this variable in JavaScript, so we define a JavaScript cell to do that. However, we cannot simply access the variable declared in a Python cell without first converting it using one of the magics given above. Try it anyway, and see if you understand the resulting error.

```sos
%use javascript
// try to modify the python variable without converting it to JavaScript first
my_var = my_var + 5
```

So let's do it properly: We define a JavaScript cell, where we import the variable `my_var` from python using `%get my_var `, then modify it and finally convert it back to python using `%put my_var `. 

```sos
%use javascript
%get my_var 
%put my_var 
// modify the imported variable (with the same name)
my_var = my_var + 5;
```

If we now print the value of `my_var` in Python, it should be 5.

```sos
%use sos
# Print the value of the python variable my_var
print(f'The value of my_var is {my_var}')
```

Note that, in the previous-to-last cell, the command `%put my_var ` only makes its effect **after** the whole cell terminates. In this case, it takes effect after the line `my_var = my_var + 5;` runs. Similarly, `%get my_var ` runs **before** the cell is run.

# <a name="2_python_and_js"></a> 2. Python and JavaScript

Now: why do we need to make everything so complicated and use two different programming languages? We can assure you it's not only to annoy you, and not only because we like the challenge. The real reason is because Python and JavaScript excel at two very distinct tasks. On one hand, **Python** is a high-level language with many great and popular libraries to **apply image processing algorithms** and visualize images. On the other hand, **JavaScript** is a fairly low-level programming language, and it is great for **implementing image processing algrithms** by accessing each pixel in an image. 

The purpose of this course is not only for you to understand how to apply the different image processing methods to images, for which we will mostly use Python, but you should also learn how methods work in detail, pixel by pixel, for which we will mostly use JavaScript. In the upcoming labs you will usually need to implement every method **once** in JavaScript to show that you really understand the concept, and then we will show you how to apply the same method in Python, using existing image processing libraries. Remember that you don't need to worry about not knowing any of the two languages! You will be provided with examples of everything that you need to do. Of course, you always have the option to contact your TAs or check [Stack Overflow](https://stackoverflow.com/), where you can find the answers to thousands of previously asked questions.

## <a name="2A"></a> 2.A. Basic Python Coding
Declaring a variable in python is very simple. We have three basic types of variables:
```python
# Numbers
my_var = 2
my_var = 5.2
# Strings
my_var = 'Either use single quotes'
my_var = "or use double quotes"
# Boolean
my_var = True
my_var = False
```
Every variable can hold every type of data, so we don't need to specify the type at initialization.

We can have lists of several items:
```python
# List of numbers
my_list = [2.5, 7, 3.3]
# List of strings
my_list = ['E.T.', 'phone', 'home']
# List of booleans
my_list = [True, True, False]
```
Lists are 0-indexed and can be accessed with square brackets:
```python
# Set the 3rd element of my_list to 5
my_list[2] = 5
# Retrieve the 2nd element of my_list
new_var = my_list[1]
```
Similarly we can make tuples, which are just like lists, but their values cannot be modified after they have been initialized:
```python
# This is a tuple of three numbers
my_var = (3, 8.6, 4)
# Access to a tuple is similar to lists
new_var = my_var[1]
# But we cannot write to a tuple!
```
For numbers we have the following basic operators:
```python
# Define the two variables
a = 5
b = 2
# Addition
a + b
>> 7
# Subtraction
a - b
>> 3
# Multiplication
a * b
>> 10
# Division
a / b
>> 2.5
# Integer division
a // b
>> 2
# Modulo (remainder of integer division)
a % b
>> 1
# Power
a ** b
>> 25
```
We can have conditional statements
```python
if a < b:
    print('a is smaller than b')
else:
    print('a is bigger than b')
```
<div class="alert alert-info">

**Note:** Indentation is important in python. Writing a conditional statement, a loop, a function, etc., without an indent (`Tab`) raises an error.
</div>

with the following comparison operators:
```python
<  # Smaller than
<= # Smaller or equal
>  # Bigger than
>= # Bigger or equal
== # Equal
!= # Not equal
```
A `for` loop is declared as:
```python
# Note: range(n) creates numbers from 0 to n-1
for i in range(10):
    # Another for loop inside the for loop - we call this 'nested for loops'
    for j in range(10):
        # We can also have formatted strings indicated by a preceeding 'f', where the number j is inserted into the string
        print(f'Loop iteration number {j}')
```

Similarly a `while` loop:
```python
i = 0
while i < 10:
    print(f'Loop iteration number {i}')
    i = i + 1
```
And finally, a function can be defined using the `def` keyword:
```python
# A function that adds two numbers and returns the result
def addition(a, b):
    # Add the numbers
    result = a + b
    # Return the result
    return result

# Run the function
addition(5, 10)
>> 15
```

Experiment with Python in the cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above.

```sos
%use sos

```

### <a name="2Aa"></a> 2.A.a. NumPy

A very important addition to Python is the [NumPy](https://NumPy.org/) library, which provides multidimensional matrix/array processing tools. Most image processing libraries in Python use NumPy arrays to represent images. We will do the same, and therefore it is important you know how to use them.

We usually import the NumPy library as `np` so we don't have to type `NumPy` every time, as well as to make the code more readable. Run the cell below to import NumPy.

```sos
%use sos
# Import the NumPy library
import numpy as np
```

To create a new NumPy array filled with zeros we can use `np.zeros`. It requires the size of the array we want to get as a tuple `(height, width)`. Run the cell below of an example.

Optionally, we can specify the [type](https://NumPy.org/devdocs/user/basics.types.html) of the array. For example, one can use 8-bit unsigned integers (`np.uint8`), which is convenient when working with 8-bit images, as the resulting array will occupy less memory than with other types, such as `np.float32` or `np.float64`. For most of the labs, you will not need to mind the type of your variables: it adapts to represent what you are storing in each instance.

```sos
%use sos
# A NumPy array of size 5x10 with type unsigned 8-bit integer
my_arr = np.zeros((5,10), dtype=np.uint8)
# Print the result
print(my_arr)
```

NumPy arrays are accessed like lists, but with an index for every dimension. Before running the example below, try to guess where the `1` will appear in `my_arr`, i.e., what are $Y$ and $X$ in the comments.

<div class="alert alert-warning">
    
**Beware:** In NumPy arrays, the first index specifies the row (y-axis location), not the x-axis location, while the second index specifies the column (x-axis location), not the y-axis location.
</div>

```sos
%use sos
# Set the value at the Yth row, Xth column to 1
my_arr[3, 7] = 1
print(my_arr)
```

We can also use *slicing* (:) to determine whole regions of the array to read or assign. Before running the example, guess where the `1`s will appear, i.e., what are $Y_1,Y_2$ and $X_1,X_2$ in the comments.

```sos
%use sos
# Insert a rectangle with the upper left corner at (Y_1,X_1)
# and the lower right corner at (Y_2,X_2)
my_arr[:3, 2:6] = 1
print(my_arr)
```

As you can see in the code, in slicing, we define an interval of numbers by `<lower_limit>:<upper_limit + 1>`. If one of the limits is not specified, it is interpreted as _from the beginning_ or _to the end_. 

Using slicing we can also extract patches of an array to create new arrays. Run the example below and see its effect.

```sos
%use sos
# Let's extract the rectangle that we created
my_rect = my_arr[:3, 2:6]
# Note: '\n' creates a new line when printing
print(f'Original array: \n{my_arr}\n\nExtracted patch: \n{my_rect}')
```

Using NumPy arrays we can easily perform the basic operators on all elements of one or several arrays (element-wise operations). Run the cell below to define a random array to work with.

```sos
%use sos
# First we define an array of size 5x5 filled with random integers from 0 to 10
np.random.seed(uid)
rand_array = np.random.randint(low=0, high=10, size=(5,5))
print(f'Initial array:\n{rand_array}')
```

We can perform any operation with a scalar for all elements of the array at the same time by using the array as if it was a variable with a single number. Run the example below and make sure you understand its results.

```sos
%use sos
# Add 2 to all elements
rand_array = rand_array + 2
print(f'rand_array + 2:\n{rand_array}\n')
# Divide all elements of the resulting array by 2
rand_array = rand_array / 2
print(f'rand_array / 2:\n{rand_array}')
```

Similarly, we can use expressions with multiple arrays of the same size and any of the basic operators: the calculation will be made on the co-positioned elements of the different arrays. Run the example below and make sure you understand its results.

```sos
%use sos
# Let's initialize two new random arrays
rand_array_1 = np.random.randint(low=0, high=10, size=(5,5))
rand_array_2 = np.random.randint(low=0, high=10, size=(5,5))
print(f'Array 1:\n{rand_array_1}\n\nArray 2:\n{rand_array_2}')
# We can for example add them together
added_arrays = rand_array_1 + rand_array_2
print(f'\nArray 1 + Array 2:\n{added_arrays}')
# Or multiply them
multiplied_arrays = rand_array_1 * rand_array_2
print(f'\nArray 1 * Array 2:\n{multiplied_arrays}')
```

We can also perform _advanced indexing_ of NumPy arrays: we can provide a boolean array as the index and extract/modify only the values that are specified as `True` by the boolean array. An application of this form of indexing could be to take the multiplied array from the previous cell and set all its elements that are larger than $25$ to $100$. Run the example below and make sure you understand its results.

```sos
%use sos
# First print the multiplied array as a reference
print(f'Multiplied array:\n{multiplied_arrays}')
# Create the boolean array that contains true at every location, where the multiplied array has a value bigger than 25
boolean_array = multiplied_arrays > 25
print(f'\nBoolean array:\t(Multiplied array > 25)\n{boolean_array}')
# Now we can use this array to index the multiplied array and set all elements indicated by True to 100
multiplied_arrays[boolean_array] = 100
print(f'\nMultiplied array with all elements > 25 set to 100:\n{multiplied_arrays}')
```

As you can see, NumPy provides extremely useful tools to work on images represented as NumPy arrays. During the labs, you will get to know even more useful functions in the NumPy library, and they will make your life much easier.

### <a name="2Ab"></a> 2.A.b Matplotlib

In order to use the incredible tools NumPy provides on images, as well as to _see_ their effect, we need to be able to 1) turn images into NumPy arrays and 2) visualize the NumPy arrays as images. That's where [Matplotlib](https://matplotlib.org/) comes into play, specifically the pyplot sublibrary: it provides all the necessary tools to import and visualize images.

We usually import `matplotlib.pyplot` as `plt` to make the code more readable. Run the cell below to import it.

```sos
%use sos
# Set the display to be inside the notebook
%matplotlib widget
# Import pyplot
import matplotlib.pyplot as plt
```

The only functionality you really need to know is how to load images using `plt.imread`:

```sos
%use sos
# Load an example image located at 'images/example_image.png'
example_img = plt.imread('images/example_image.png')
# Then, we can check the image is really a NumPy array
type(example_img)
```

Just so you can actually see the image you loaded, we will display it using `plt.imshow`. In the labs, you will never use this function. Instead, we provide the [IPLabViewer](#4_viewer) class, which adds extra functionality on top of `plt.imshow`, is more convenient, and is easier to use.

```sos
%use sos
# When using plt.imshow we first need to declare a new figure
plt.figure( )
# Then we add the image to the figure
plt.imshow(example_img)
# Give it a title
plt.title(f'Example Image - SCIPER: {uid}')
# And finally we make sure it is displayed on the screen
plt.show()
```

Experiment with Python, NumPy and Matplotlib in the cell below. Learn to do the simple things that come to mind and take the chance to explore all of the options above.

```sos
%use sos
```

## <a name="2B"></a> 2.B. Basic JavaScript Coding
Declaring a variable in JavaScript uses the `var` statement. Similarly to what we saw with [Python](#2A) we have three basic types of variables:
```javascript
// numbers
var my_var = 2;
var my_var = 5.2;
// strings
var my_var = 'Either use single quotes';
var my_var = "or use double quotes";
// boolean
var my_var = true;
var my_var = false;
```
Every variable can hold every type of data, so we don't need to specify the type at initialization.
<div class="alert alert-info">

**Note:** In JavaScript, every line should be terminated with a semicolon `;`.
</div>

We can have lists of several items:
```javascript
// list of numbers
var my_list = [2.5, 7, 3.3];
// list of strings
var my_list = ['E.T.', 'phone', 'home'];
// list of booleans
var my_list = [True, True, False];
```
Lists are 0-indexed and can be accessed with square brackets:
```javascript
// set the 3rd element of my_list to 5
var my_list[2] = 5;
// retrieve the 2nd element of my_list
var new_var = my_list[1];
```
Contrary to Python, tuples do not exist in JavaScript.

For numbers we have the following basic operators:
```javascript
// define the two variables
var a = 5;
var b = 2;
// addition
a + b;
>> 7
// subtraction
a - b;
>> 3
// multiplication
a * b;
>> 10
// division
a / b;
>> 2.5
// integer division can be emulated through
parseInt(a / b);
>> 2
// modulo (remainder of the integer division)
a % b;
>> 1
// power
a ** b;
>> 25
```
We can have conditional statements
```javascript
if(a < b){
    console.log('a is smaller than b');
}else{
    console.log('a is bigger than b');
}
```
<div class="alert alert-info">

**Note:** Indentation is **not** important in JavaScript but curly brackets `{}` are. When writing a conditional statement, a loop, a function, etc., one uses curly brackets to define the respective beginning and end.
</div>

with the following comparison operators:
```javascript
<  // smaller than
<= // smaller or equal
>  // bigger than
>= // bigger or equal
== // equal
!= // not equal
```
A `for` loop is declared as:
```javascript
for(var i = 0; i < 10; i++){
    // another for loop inside the for loop
    for(var j = 0; j < 10; j++){
        // we can concatenate strings and numbers using the + operator
        console.log('Loop iteration number ' + j);
    }
}
```
Similarly a `while` loop:
```javascript
var i = 0;
while(i < 10){
    console.log('Loop iteration number ' + i);
    // increment i by one
    i++; 
}
```

And finally, a function can be defined using the `function` keyword:
```javascript
// a function that adds two numbers and returns the result
function addition(a, b){
    // Add the numbers
    var result = a + b;
    // Return the result
    return result;
}

// run the function
addition(5, 10);
>> 15
```

<div class="alert alert-info">

**Note:** In the Jupyter Notebook environment we use for the labs, writing the `var` for variable declerations is not strictly necessary, since the code works fine without it. Similarly, the semicolons at the end of every line are not needed for the code to run. That being said, it is still good practice to always add the two and we thus recommend that you do it anyway.
</div>

Experiment with JavaScript in the cells below. Learn to do the simple things that come to mind and take the chance to explore all of the options above. If you want to visualize what you try out, use the cell further below, Python, and SoS to play around. If that looks scary at this point, move on and come back some time later.

```sos
%use javascript
```

```sos
%use sos
```

## <a name="2C"></a> 2.C. Python and JavaScript speed comparison

As mentioned before, JavaScript will mainly be used to implement pixel-level algorithms, which require nested `for` loops. To illustrate why this is necessary, we'll compare the time it takes for both Python and JavaScript to perform a zero-padded convolution implemented using nested `for` loops.

The two-dimensional convolution between two discrete images $h[m,n]$ and $f[m,n]$ is given by

$$(h \ast f)[k,l] = \sum_{m \in \mathbb{Z}}\sum_{n \in \mathbb{Z}}f[m,n]h[k-m,l-n]$$
Don't worry, you don't need to know how that works right now, you'll see this in Chapter _3.2_ of the course. But it's safe to say that the convolution is a very important operation for image processing, which is why we'll use it for this demonstration.

First let's define a Python function to generate a Gaussian kernel, which we will use as one of the images for the convolution. Run the cell below to define this function and obtain and visualize the Gaussian impulse response.

```sos
%use sos
# Function that creates a Gaussian convolutional kernel based on the input parameters
def gaussian_kernel( sig, l):
    # Generate the 1D grid-points
    ax = np.linspace(-(l - 1) / 2., (l - 1) / 2., l)
    # Generate the 2D grid-points created by ax x ax
    xx, yy = np.meshgrid(ax, ax)
    # Calculate the Gaussian kernel (up to proportionality constant)
    kernel = np.exp(-0.5 * (np.square(xx) + np.square(yy)) / np.square(sig))
    # Normalize the kernel to sum 1 and return it
    return kernel / np.sum(kernel)

# Get a Gaussian kernel
ker = gaussian_kernel(sig=3, l=5)
# Show the result
plt.figure()
plt.imshow(ker, cmap='gray')
plt.title(f'Gaussian kernel - SCIPER: {uid}')
plt.show()
```

From the implementation above you can learn how to use NumPy functions such as `np.linspace`, `np.meshgrid`, `np.exp`, `np.square` and `np.sum`. Try them out at the end of the Python section if you are curious! Otherwise, don't worry, this example will always be here for you.

Now that we have the kernel ready, let's define the JavaScript convolution function. Run the cell below to define it.

```sos
%use javascript
// function that convolves img with kernel
function convolution(img, kernel){
    // initialize the output image as a copy of the input image
    var smoothed_img = [...img];
    // compute the kernel offset
    const k_offset = Math.floor(    kernel.length / 2 );
    const l_offset = Math.floor( kernel[0].length / 2 );
    // perform zero-padded convolution
    for(i = 0; i < img.length; i++){
        for(j = 0; j < img[0].length; j++){
            // compute output for pixel i, j. 
            // zero-padding: extend the image with zeros beyond its boundaries.
            smoothed_img[i][j] = 0;
            for(k = 0; k < kernel.length; k++){
                for(l = 0; l < kernel[0].length; l++){
                    // compute evaluation points to enforce boundary
                    var x = k - k_offset + i;
                    var y = l - l_offset + j;
                    // increment if appropriate (if outside the image, don't increment, there are zeros)
                    if((x < img.length) && (x > 0) && (y < img[0].length) && (y > 0)){
                        smoothed_img[i][j] += img[x][y] * kernel[k][l]; 
                    } 
                }
            }
        }
    }
    // return the output image
    return smoothed_img;
}
```

From the implementation above, you can learn useful JavaScript functions like `Math.floor`. Nonetheless, some of the tricks we had to use will not be necessary for you. For example, the lines
```javascript
var smoothed_img = [...img];
k < kernel.length;
l < kernel[0].length;
```
will be much more intuitive to implement using our class [IPLabImageAccess](#2_ImageAccess) below, so do not worry about them. 

Now, let's define the Python convolution function. Run the cell below to define it.

```sos
%use sos
# Function that convolves img with kernel
def convolution(img, kernel):
    # Initialize output image as an empty image
    smoothed_img = np.zeros(img.shape)
    # Compute the kernel offset
    k_offset = kernel.shape[0] // 2
    l_offset = kernel.shape[1] // 2
    # Perform zero-padded convolution
    for i in range(img.shape[0]):
        for j in range(img.shape[1]):
            # Compute output for pixel i, j. 
            # Zero-padding: extend the image with zeros beyond its boundaries.
            for k in range(kernel.shape[0]):
                for l in range(kernel.shape[1]):
                    # Compute evaluation points to enforce boundary
                    x = k - k_offset + i
                    y = l - l_offset + j
                    # Increment if appropriate (if outside the image, don't increment, there are zeros)
                    if (x < img.shape[0]) and (x > 0) and (y < img.shape[1]) and (y > 0):
                        smoothed_img[i, j] += img[x, y] * kernel[k, l]
    # Return the output image
    return smoothed_img
```

From the implementation above, you can learn useful NumPy functions like `np.zeros`, or methods like `img.shape`, when `img` is a NumPy array. As always, feel free to try them out now at the end of the Python section!

Now, we load an image and transfer it to JavaScript, together with the Gaussian kernel we defined before in Python. In this way, we will be able to test the JavaScript function we implemented. Run the cell below to load the image, display it, and transfer the image and the kernel to Javascript.

```sos
%use sos
%put hrct --to javascript
%put ker --to javascript
# Load the hrct image
hrct = plt.imread('images/hrct.tif')
# Display it
plt.figure()
plt.imshow(hrct)
plt.title(f'HRCT - SCIPER: {uid}')
plt.show()
```

Now, we are ready to run the convolution in JavaScript and measure the time it takes. Run the cell below to do it.

```sos
%use javascript
%put convolved_hrct_js 
%put js_time 
// get the starting time in miliseconds
var start = new Date();

// run the convolution
var convolved_hrct_js = convolution(hrct, ker);

// get the ending time in miliseconds
var end = new Date();

// calculate the elapsed time in miliseconds
var js_time = end - start

// display the elapsed time
console.log('The JavaScript convolution took ' + js_time + ' miliseconds to run.');
```

Now, we do the same for the implementation of the convolution in Python. Run the cell below to do it.
<div class="alert alert-danger">

**Spoiler alert:** You can go grab a drink while this is running... As you may have noticed, in Jupyter notebooks a cell is still running if inside the square brackets in the top left corner of the cell you see an asterisk `[*]` instead of a number.
</div>

<div class="alert alert-info">

**Note:** To measure the time in Python we need to import the `time` module.
</div>

```sos
%use sos
import time

# Get starting time in seconds
start = time.time()

# Run the convolution
convolved_hrct_jpython = convolution(hrct, ker)

# Get the ending time seconds
end = time.time()

# Calculate the elapsed time in miliseconds
python_time = round((end - start)*1000)

# Display the elapsed time
print(f'The Python convolution took {python_time} miliseconds to run.');
print(f'That makes JavaScript around {round(python_time/js_time)} times faster than Python for this implementation.')
```

Hopefully, you can understand from this illustration that standard Python is not made for pixel-level algorithms with nested `for` loops. This justifies the use of JavaScript for such implementations. Just so you can actually see the result of the function we just ran, run the next cell to see it.

```sos
%use sos
plt.figure()
plt.imshow(convolved_hrct_jpython)
plt.title(f'HRCT convolved with a Gaussian kernel - SCIPER: {uid}')
plt.show()
```

# <a name="2_ImageAccess"></a> 3. IPLabImageAccess (2 Points)
While JavaScript offers us the big upside of fast `for` loops, it also comes with a downside. It does not offer any convenient array/image processing library like Python (with NumPy or OpenCV, which you'll get to know during the labs). To make things easier, we created our own JavaScript class, `IPLabImageAccess`, which facilitates the handling of images in JavaScript. Every `IPLabImageAccess` object is a container for an image and provides easy-to-use methods for pixel access that take care of boundary conditions. This will help you to complete the tasks in JavaScript. Here, we will show you the basics of how this class can be used.

As we do with external libraries, we will import this class under a simpler name to make the code more readable. In all upcoming labs, we will import the `IPLabImageAccess` class as `Image`. Run the JavaScript cell below to do it.

```sos
%use javascript
const Image = require('./lib/IPLabImageAccess.js')
```

Once imported, we can start using it. First of all, let's look at the different ways we can create a new `Image` object.

* An `Image` object can be created with
```javascript
var img = Image(height, width);
```
This will create a new `Image` object of dimensions `height` $\times$ `width`. <br><br>

* Optionally, `height` and `width` can be provided as a single array `[height, width]`. This makes it easy to create an image with the same dimensions as an existing image by using the `shape` method, for example
```javascript
var img = Image(an_existing_image.shape());
```

* By default, the image is initialized to $0$ for every pixel, but we can choose an alternative initialization value by specifying the `init_value` parameter while creating a new `Image` object, for example
```javascript
var img = Image(height, width, init_value=255);
```

* An `Image` object can also be created from an existing JavaScript array by simply passing it as a single argument
```javascript
var img = Image(an_existing_array);
```
This is useful, for example, to create an `Image` object from a Python NumPy array converted to Javascript.

We now look at how to use the basic functionalities in the class. The most important operations you're gonna need are the `img.setPixel` and `img.getPixel` methods, which do exactly as their name suggests when `img` is an `Image` object.
* The `img.setPixel(x, y, value)` method can be used to set the pixel at location `(x,y)` to a specific value `value`:
```javascript
// here we set the pixel (3,5) of img to 255
img.setPixel(3, 5, 255);
```
* The `img.getPixel(x, y)` method can be used to retrieve the value of the pixel at location `(x,y)`:
```javascript
// here we retreive the pixel (3,5) of img
var pixel = img.getPixel(3, 5);
```
By default, the `setPixel` and `getPixel` methods will use _mirroring_ boundary conditions to provide values for pixels outside the image range. Do not worry if you do not understand this yet, it will be clearer as you progress through the course.<br><br>

* The height and width of an `Image` object can be retreived either using using the `img.shape()` method, which returns a list `[height, width]` or, more conveniently, with `img.nx` and `img.ny`.
```javascript
// retreive height and width of the image as a list
height_and_width_list = img.shape();
// get only the width
width = img.nx;
// get only the height
height = img.ny;
```
This second option is very useful when setting the limits in `for` loops. <br><br>

* We can extract the neighbourhood of size `width` $\times$ `height` around the pixel at location `(x,y)`, using the `img.getNbh(x, y, width, height)` method:
```javascript
// extracting the 3x3 neighbourhood around pixel (12,7)
nbh = img.getNbh(12, 7, 3, 3);
```
By default, the method will use _mirroring_ boundary conditions, just like the `getPixel` and `setPixel` methods.

* Finally, if we need to convert an `Image` object to a simple JavaScript array, we can do it with the `img.toArray()` method, for example
```javascript
// convert the Image object to a JavaScript array
img_array = img.toArray();
```
This will be necessary, for example, to convert the image to Python.

We are sorry you just had to read so much! For **2 points**, your job now is to create a new `Image` object called `img`, of size $50 \times 50$, and set the pixels at locations $(x=10, y=20)$, $(x=15, y=30)$ and $(x=40, y=7)$ to $255$.

<!-- grade_id: cell-779a0057e28178dd, points: 0, grade: false, solution: true -->
```sos
%use javascript
%put img 

// create a new Image object named img of size 50 x 50
# YOUR CODE HERE

// set the pixel at location (x=10,y=20) to 255
# YOUR CODE HERE
// set the pixel at location (x=15,y=30) to 255
# YOUR CODE HERE
// set the pixel at location (x=40,y=7) to 255
# YOUR CODE HERE
```

Run the cell below to (partially) check if you did the right thing.

<!-- grade_id: cell-e5d66a3787cd6867, points: 2, grade: true, solution: false -->
```sos
%use javascript
if( (img.nx != 50) || (img.ny != 50) ){
    throw new Error("The image size is not correct.");
}
if( img.getPixel(7,40) == 255 ){
    throw new Error("The pixel at location x=7, y=40 should have value 0.")
}

```

# <a name="4_viewer"></a> 4. IPLabViewer

As you saw in [Section 2.A.a: NumPy](#2Aa), the main reason that Python is so powerful is because of its libraries. To ease image visualization and manipulation, we have developed a library containing an **IPLabViewer** class, which you will be using throughout the course. This section explains its functioning and shows some examples and use cases of the **IPLabViewer**. Before using the class, be sure to check the documentation. You can call `help(viewer)` or `viewer?` in a cell, but we recommend you to check its extensive [wiki](https://github.com/poldap/IPLABs2020/wiki/Python-IPLabViewer()-class). 

In the following cell, we import the required libraries. For ease of use, in the labs throughout the course we will import it as `viewer`. Moreover, we also load some of the images that you will be using throughout this lab.

The IPLabViewer runs without explicitly importing any external libraries (all are imported in the class' `.py` file). However, we will import the library Ipywidgets ([see documentation](https://ipywidgets.readthedocs.io/en/latest/)) to allow interactive visualization. You will even use it to add your own functions to the viewer's interface! 

The magic command `%matplotlib widget` configures Matplotlib to use the Ipywidgets library to render dynamic plots. It allows us, for example, to zoom into images (button with a small square), and allows for proper functioning of the widgets. The IPLabViewer class **should not** be used outside the dynamic environment of Matplotlib + Ipywidgets. 

Run the cell below to import all necessary libraries and load the images we will work with.

```sos
%use sos
# Configure plotting as dynamic
%matplotlib widget
# Import required packages for this exercise
import matplotlib.pyplot as plt
import ipywidgets as widgets
# Import IPLabViewer
import sys  
sys.path.insert(0, 'lib')
from iplabs import IPLabViewer as viewer

# Import images
car   = plt.imread('images/car_pad.tif')
hrct  = plt.imread('images/hrct.tif')
hands = plt.imread('images/hands.tif')
plate = plt.imread('images/plate.tif')
boats = plt.imread('images/boats.tif')
```

## <a name="instantiation_viewer"></a> 4.A.: Creation of a viewer

The next cell will illustrate many of the functionalities of IPLabViewer, as well as its customization capabilities. Run it and explore its results while reading the explanatory cell below it.

```sos
%use sos

# Declare image list
first_list = [hands, plate, car, hrct]
# Declare titles (only two)
title_list = ['Hands', 'Plate']
# Call the image viewer with the image list as first argument. The rest of the arguments are optional
first_viewer = viewer(first_list, title = title_list, colorbar = True, widgets = True, 
                      hist = False, axis = True, cmap = 'spring')
```

To create an IPLabViewer object, you simply call the function `viewer`. You can do that without assigning a variable name, but it is best to do so. Ideally, the variable should be self-explanatory and end with the word *viewer* to avoid confusion (for example, `first_viewer` as above).

* Images: The same viewer can load many images at once. In this example, it was created with $4$ images. By default, you can switch between the images using the buttons `Prev` and `Next` on the widgets panel.
* Titles: You can give the viewer titles for each image, which will be shown immediately above it. If you do not provide a title for some images, the variable name will be shown by default. In this example, only $2$ were given.
* Customization: Several of the class' functionalities (e.g., `widgets`, `hist`, `axis`, `colorbar`) can be enabled from the beginning. A user can later turn them off using the widgets or [calling methods](#prog_customization) of the object. 
* Colormap: The colormap to be used can be set from the start with the parameter `cmap = 'colormap'`, and dynamically changed using `Options` and choosing one of the supported colormaps in the drop-down menu. 
* Statistics: On the lower part of the widget panel, the image statistics of the area of the image being displayed are shown. If you use the button with a little square at the left of the images to zoom into a selected region, you will see the statistics are updated in real time. You can also use the button with crossed double-arrows to pan through the image at a specific level of zoom. You can always use the button *Reset Original View* to see the whole image again.

<div class="alert alert-warning">
    
<b>Beware: </b> The statistics displayed in the <b>IPLabViewer</b> objects are rounded to the second decimal. If you need an exact statistic, use the methods `np.mean()` or `np.std()` on the images (NumPy arrays) directly.
</div>

* Display mode: There are several options to display a list of images
    * Default display: Display only one image at a time, and use the buttons *Next* and *Prev* to browse the different images. 
    * Set `subplots = [m, n]`: Arrange the images in an $m \times n$ grid.
    * Set `column_view = True`: Display all the images in one column. This option has priority over `subplots = [m, n]`.

## <a name="widgets_viewer"></a> 4.B.: Using the widgets

If you just want a quick visualization, you can instantiate the IPLabViewer with the image you want to display (or list of images). All the parameters will be set to their default values (`cmap = 'gray'`, vertical display, all the features set to off), but you can use the widgets to change this. Run the next cell to see all the default values. Go through all the widgets and explore their options (click on the button *Show Widgets* to activate them).

Your first view will consist of:
* A series of buttons at the left, which serve to control the dynamic environment of matplotlib. If you hover your mouse in the buttons, wou will see what each button does. You can use them to zoom to a region, to move around that region, and to save an image in png format.
* The figure holding the images (one figure is holding all the images). You can change the size of a figure using the small gray triangle (bottom-right corner). If you hover your cursor over the triangle, you will see it change to a two-sided arrow. By clicking on it you will be able to adjust the figure size.
* A button with the legend *Show Widgets*. Clicking it will take you to the widget main menu. 

In the widgets main menu you will see the following widgets:
* *Brightness and Contrast* Slider: In this menu you will be able to change the color scaling of the image through a slider, given in percentages of original.
* *Show Histogram* Button: Show or hide the histograms of the images.
* *Options* Button: It will take you to the options menu, where you will be able to show or hide the axis, the colorbar, and change the colormap.
* *Reset Origianl View* Button: If you clicking after doing any operations, all will be reverted (the colormap will be set to the original one, it will hide the colorbar and the axis, reset the color scaling and the zoom).
* Stats Textbox: Here you will see the mean, the standard deviation, the range of values and the size of the image. Mean, standard deviation and range of values will be updated when you zoom to a region. 

Run the next cell and interact with the viewer.

```sos
%use sos

# In order to keep the memory load on Noto low, 
# we will generally close all figures before continuing.
# This means that the previous viewers are no longer interactive!
plt.close('all')

# Declare image list
second_list = [hrct, hands]
# Call the image viewer with minimal arguments
second_viewer = viewer(second_list)
```

## <a name="usr_widgets_viewer"></a>4.C.: User defined widgets

**IPLabViewer** has several cool functionalities. But what if you wanna contribute yourself? **IPLabViewer** allows you to create a function in a Jupyter notebook, and apply it simultaneously to all images within your **IPLabViewer** object for a set of slider values. The function you create will take as parameters:
* an image (`NumPy array`), and 
* one or more numerical values.

Your function is then supposed to apply an operation on the image that depends on these value(s). This is very useful if you want to see the effect of the value(s) you choose on one or several images. Usually, you would have to manually run the same process several times, and visualize the results each time. With **IPLabViewer**, you can simply use slider(s) within the viewer.

In the next cells, you will see a very basic example. We will apply a pixelwise operation on the image `car`: All the pixels with a value below a treshold (given in $\%$) will be set to the *maximum* value. We will guide you through all the process.

First, we define the thresholding function `your_function`. Run the next cell to define it.

```sos
%use sos
# Importing NumPy
import numpy as np
# Define your function
def your_function(image, threshold):    
    # We make a copy of the original, where we will apply our threshold.
    output = np.copy(image)
    # For greater flexibility, we will get the value as a percentage of the maximum
    value = threshold*0.01*np.amax(image)
    # Apply threshold
    output[image < value] = np.amax(image)
    return output
```

Besides defining your function, you need to define the corresponding slider(s), a button to run `your_function` with the sliders' values, and the specific function that will run `your_function` (a.k.a., the slider's _callback_ function).

Although this is rather abstract, the example will make it quite clear. We will use an integer slider in the range $[0,100]$ (`widgets.IntSlider`), a button that clearly specifies what our function does (`widgets.Button`), and a callback function `callback` that calls `your_function` with the slider's value as threshold. Run the next cell, click the button *Extra Widgets*, and try different values. Remember to click on `Apply Threshold`!

```sos
%use sos

plt.close('all')

# Declare slider
threshold_slider = widgets.IntSlider(value = 0,min = 0, max = 100, step = 1, 
                                     description = 'Threshold', continuous_update = False )

# Declare button with meaningful description
activation_button = widgets.Button(description = 'Apply Threshold')

# Declare callback
def callback(img):
    # Get slider value
    threshold = threshold_slider.value
    # Call your function
    output = your_function(img, threshold)
    return output

# Call viewer, passing the widget and callback separately as lists
threshold_viewer = viewer(car,  new_widgets = [threshold_slider, activation_button], callbacks = [callback], widgets = True)
```

If you're interested in this and similar topics, we invite you to check the documentation for Ipywidgets [here](https://ipywidgets.readthedocs.io/en/stable/).

## <a name="prog_customization"></a> 4.D.: Programmatic customization

Most options can be modified both through widgets and _programmatically_ (in code). Run the following cells. In the first one, we will create the object `car_viewer`, which displays the image `car`. By running the next cells, one by one, we will produce changes on the `car_viewer` without interacting with it directly. 

```sos
%use sos
plt.close('all')
car_viewer = viewer(car)
```

```sos
%use sos

car_viewer.set_axis(axis = True)
```

```sos
%use sos

car_viewer.set_widgets(widgets = True)
```

```sos
%use sos

car_viewer.set_colormap(colormap = 'viridis')
```

```sos
%use sos

car_viewer.show_histogram(hist = True)
```

```sos
%use sos

car_viewer.set_axis(axis = False)
```

```sos
%use sos

car_viewer.set_colorbar(colorbar = True)
```

## <a name="ex_viewer"></a> Try it yourself!

### Creating an IPLabViewer (1 Point)
Now that you have explored the capabilities of the **IPLabViewer** class, create in the next cell a viewer object of the class with the following characteristics (do it with code, on one or several lines):
* The first image should be `hrct` and the second `plate`.
* The titles are set to default
* The colormap is set to 'inferno'
* There is a colorbar 
* The axis are set to `False`
* Only a single image at a time is displayed

Store your object in the variable `exercise_viewer`

<!-- grade_id: cell-7c287f72627b42ee, points: 0, grade: false, solution: true -->
```sos
%use sos
plt.close('all')
exercise_viewer = None
# YOUR CODE HERE
```

If you want to check the state of your **IPLabViewer** object, you can check the documentation and explore the attributes. See the two examples in the next cell. Both should show that there are indeed two images.

<!-- grade_id: cell-18c524bc087e7f47, points: 1, grade: true, solution: false -->
```sos
%use sos

print(len(exercise_viewer.image_list))
print(exercise_viewer.number_images)
```

### Threshold Operation (3 Point)

Now you will code a function and add a slider to apply this function to the image (see Section [4.C.: User Defined Widgets](#usr_widgets_viewer)).

For **1 point**, your first task is to define a function `range_function(img, rmin, rmax)` that will only keep pixel values inside a range of values $[r_{\mathrm{min}}, r_{\mathrm{max}}]$ ($r_{\mathrm{min}}$ and $r_{\mathrm{max}}$ included), while the values outside this range will be set to $0$.
You will apply it on the image `car`, so `rmin` and `rmax` should be between $0$ and $255$. Define this function in the following cell. 

<div class="alert alert-info">
  
<b>Note: </b> Use boolean indexing of the NumPy array `img`, as in the thresholding example, to perform the necessary operations.
</div>

<!-- grade_id: cell-f7d3dd947484c706, points: 0, grade: false, solution: true -->
```sos
%use sos

# Define your function
def range_function(img, rmin, rmax):    
    # Best practices are creating a copy. You don't wanna mess the original.
    output = np.copy(img)
    # set the pixels of output that are smaller than rmin or bigger than rmax to 0
    # YOUR CODE HERE
    return output
```

You can test your function with the following cell. In it, we define a linear array `test_array` that is formed by the integers from 0 to 255 in ascending order, using the method [NumPy.arange](https://NumPy.org/doc/stable/reference/generated/NumPy.arange.html). Then, we apply your function with the range $[100, 199]$ to it. There should be exactly $100$ elements that were not set to $0$. If the cell runs smoothly, your implementation is probably correct.
<div class="alert alert-info">
    
**Note:** The function `np.count_nonzero(arr)` counts the number of elements in `arr` that are not zero.
</div>

<!-- grade_id: cell-d32bcff0667c88a9, points: 1, grade: true, solution: false -->
```sos
%use sos

test_array = np.arange(0, 256)
assert np.count_nonzero(range_function(test_array, 100, 199)) == 100

```

For **1 point**, declare the necessary widgets to control this function (slider(s) and a button, as well as that button's callback function). For example, you can use two `IntSlider` widgets like the one we used in Section [4.C.: User Defined Widgets](#usr_widgets_viewer), and get the value of these sliders in a button's callback function. Another option is to use just one `IntRangeSlider` ([see doc](https://ipywidgets.readthedocs.io/en/latest/examples/Widget%20List.html)). When you get the value of this slider (by calling `IntRangeSlider.value`), you get a list with two elements corresponding to the minimum and the maximum value of the range. You can use these values to call your function (`range_slider.value[0]` and `range_slider.value[1]` if `range_slider` has been defined as an `IntRangeSlider` widget). Use the next cell to declare your widgets and create an **IPLabViewer**. Store your viewer in the variable `range_viewer`, and call your activation function `range_callback`.

<!-- grade_id: cell-a4cde69b86b888bf, points: 0, grade: false, solution: true -->
```sos
%use sos
# Declare the slider 
range_slider = None
# YOUR CODE HERE
# Declare the button
range_button = None
# YOUR CODE HERE
# Declare the callback
def range_callback(img):
    # Get slider value
    # YOUR CODE HERE
    # Call your function, e.g., output = range_function(img, ...)
    # YOUR CODE HERE
    return output

# Call viewer, passing the widget and callback separately as lists, e.g., range_viewer = viewer(car, ...)
plt.close('all')
# YOUR CODE HERE
```

<!-- grade_id: cell-fa45b07997e9a3be, points: 1, grade: true, solution: false -->
```sos
%use sos

```

Now, for **1 point**, and to finish the section on the **IPLabViewer** class, answer the following question using the viewer that you just defined. Note: You do not need to use your newly defined slider.

In the image `car`, what is the mean value of **the carpet that the car is standing on**? Hint: Zoom to several regions that are not too close to the edge nor to the car.

1. Around 40
2. Around 60 
3. Around 80

In the following cell, modify the variable `answer` to your actual answer, e.g., `answer='1'`, `answer='2'`, or `answer='3'`.

```sos
%use sos
# Assign your answer to this variable
answer = 0
```

<!-- grade_id: cell-9db1b31c6c9f4ebb, points: 1, grade: true, solution: false -->
```sos
%use sos
# Let's do a sanity check - check that the answer is one of the possible choices
assert answer in [1, 2, 3], "Select one of the answers: 1, 2 or 3."

```

# <a name="5_ippython"></a> 5. Image Processing in Python (2 Points)

Part of the goal of the labs is for you to learn the most standard way to perform Image Processing tasks in Python. Therefore, after you fully understand the nature of every algorithm, and have the confidence that you can implement them yourself in JavaScript, the labs will always move towards trustworthy implementations in current image processing libraries. There are many libraries available, but throughout this course we will focus on three of the most common:

* [**`openCV`**](https://opencv.org/about/): It stands for Open Computer Vision. As the name suggests, it provides thousands of open source algorithms for computer vision applications. The fact that it supports several programming laguages (`C++`, `Python`, `Java` and `Matlab`), and that it has over 20 years of history explains the huge community of users. Thus, there is also extensive documentation and tutorials, not only in its official website, but also on several other sites.  
* [**`scipy.ndimage`**](https://www.scipy.org/) SciPy stands stands for Scientific Python, and it is one of the most widely used Python libraries (their packages include `NumPy` and `matplotlib`, which are considered standards in Python). As OpenCV, it is open-source and it has a huge community, as well as extensive documentation and tutorials. It has a myriad of applications for scientific and technical computing. In the IP courses we will focus on their module built for [Multi-dimensional Image Processing (`ndimage`)](https://docs.scipy.org/doc/scipy/reference/ndimage.html).
* [**`scikit-image`**](https://scikit-image.org/) Scikit stands for SciPy Toolkits. These toolkits are add-on packages for SciPy. They are developed and maintained individually. As SciPy, they are open-source projects. Some of these Scikits have become very well known, and have a large community of users and developers. Scikit-Image, is one of them, developed specifically for Image Processing tasks.

While there are many other very good libraries for Image Processing tasks, we will only use these three because their history and community ensure long-term support and extensive documentation and examples. Whenever we want you to use one of these libraries, we will either use a function to exemplify an operation, or we will include a link to the documentation of that particular function.

Just as with `NumPy` (`np`) or `matplotlib.pyplot` (`plt`), there are standard abbreviations to import these libraries. Run the next cell, where we will import the three of them. As you will see, the packages do not necessarily keep the name of the organization. 

```sos
%use sos
# OpenCV
import cv2 as cv
# SciPy NDImage
import scipy.ndimage as ndi
# Scikit-image
import skimage
```

# <a name="examples_iplibraries"></a> Examples

Now we will go through two very basic operations (rotation and edge detection) for you to get familiar with each of these libraries. We will show you the basic syntax, or how to call the function with the minimum necessary parameters. For the full list of parameters and options, go to the official documentation of each function. 

### Rotation in OpenCV

The basic syntax of the function is ([see full doc](https://docs.opencv.org/3.4/d2/de8/group__core__array.html#ga4ad01c0978b0ce64baa246811deeac24)):
```python
output = cv.rotate(src, rotateCode)
```

The parameters are:
* `src` (NumPy array): The source image (NumPy array), and
* `rotateCode` (flag): Specifies the rotation ([see available options](https://docs.opencv.org/3.4/d2/de8/group__core__array.html#ga6f45d55c0b1cc9d97f5353a7c8a7aac2))

As you will see, OpenCV often works with these kind of flags, so you will have to check the documentation of each particular function to know their possible values. Run the next cell to see how this function works.

```sos
%use sos

# First, we call the OpenCV function
rot_car = cv.rotate(car, cv.ROTATE_90_CLOCKWISE)

# and then we visualize
plt.close('all')
cv_rotate_viewer = viewer([car, rot_car], title = ['Original', 'Rotated'])
```

### Edge Detection in OpenCV
For a quick edge detection filter, we will use the Sobel operator *derivative* filter ([see doc](https://docs.opencv.org/3.4/d4/d86/group__imgproc__filter.html#gacea54f142e81b6758cb6f375ce782c8d)). You will go through it in detail during the course, but if you want more information now you can check [this tutorial](https://docs.opencv.org/3.4/d2/d2c/tutorial_sobel_derivatives.html).

The syntax of the function with the basic parameters is ([see doc](https://docs.opencv.org/3.4/d2/de8/group__core__array.html#ga4ad01c0978b0ce64baa246811deeac24)):
```python
output = cv.Sobel(src, ddepth, dx, dy, ksize, borderType)
```

The parameters are:
* `src` (NumPy array): source image
* `ddepth` (flag): The data type of the pixels. As a rule of thumb, you can use `cv.CV_32F`
* `dx` (int): Order of derivative in x
* `dy` (int): Order of derivative in y
* `ksize` (int): The size of the kernel. Defaults to 3 (standard Sobel filter), but OpenCV gives you the option to use extended kernels (1, 3, 5 or 7)
* `borderType` (flag): Specifies the boundary conditions ([see options](https://docs.opencv.org/2.4/modules/imgproc/doc/filtering.html?highlight=copymakeborder)). Throughout the course, we will be using `cv.BORDER_REFLECT` to mean mirroring boundary conditions when values outside the image are needed.

Run the next cell for an example on the image `plate`.

```sos
%use sos

# We apply the function on the image
edges_plate = cv.Sobel(plate, ddepth = cv.CV_32F, dx =  1, dy = 1, ksize = 3, borderType = cv.BORDER_REFLECT)

# And compare to the original
plt.close('all')
cv_sobel_viewer = viewer([plate, edges_plate], title = ['Original', 'Edges'])
```

### Rotation in SciPy NDImage

The basic syntax of the function for grayscale images is ([see doc](https://docs.scipy.org/doc/scipy-0.14.0/reference/generated/scipy.ndimage.interpolation.rotate.html)):
```python
output = scipy.ndimage.interpolation.rotate(input, angle, reshape=True, mode='reflect')
```

The parameters are:
* `input` (NumPy array): The source image, 
* `angle` (float): Rotation angle in degrees,
* `reshape` (boolean): If `True` (defualt), the output shape is adapted so that the input array is contained completely in the output.
* `mode` (string): Specify boundary conditions. We will use `'reflect'` throughout the course to mean mirroring boundary conditions when values outside the image are needed. Other options include `constant`, `nearest`, or `wrap'`.

You can already see some differences between the two libraries: `scipy.ndimage` does not work with flags, but with strings. You will still have to check the documentation to see the available options. Moreover, in this particular function, scipy allows rotation by any angle. 

Run the next cell to see the result of rotating the image `car` for an angle $\theta = 45^\circ$. In the stats box of the **IPLabViewer**, look at the size of both iamges. As you can see, the resulting image has been adapted. Modify the cell, and select `mode = 'constant'`. See the difference. Experiment with different modes: if you select `mode = 'constant'`, you can add the parameter `cval` (scalar), to specify the value of the pixels outside of the image borders. You can also try to set the parameter `reshape = False` to see the result.

```sos
%use sos

# Apply operation
rot_car_ndi = ndi.interpolation.rotate(car, angle = 45, reshape=True, mode='reflect')
# Compare to original
plt.close('all')
cv_rotate_viewer = viewer([car, rot_car_ndi], title = ['Original', 'Rotated'], widgets = True)
```

### Edge Detection in SciPy NDImage

Again, we will use the Sobel filter. The basic syntax in `ndi` is ([see full doc](https://docs.scipy.org/doc/scipy/reference/generated/scipy.ndimage.sobel.html)):
```python
output = scipy.ndimage.sobel(input, axis= -1, mode='reflect')
```

The parameters are:
* `input` (NumPy array): The source image, 
* `axis` (float): axis through which to apply the 1-dimensional Sobel filter (defaults to -1, i.e., the last axis, i.e., the 2nd (1) for grayscale images)
* `mode` (string): Specify the boundary conditions (see the last example)

Again, we can see differences between the two libraries. In OpenCV you could specify different orders in each direction. In `scipy.ndimage`, the order is always 1, and you can only specify one direction. So to get an image like the one resulting from OpenCV's function, you would need to call `ndi`s funtion twice (once with `axis = 0` and once with `axis = 1`), and then get the absolute value ($\sqrt{\partial x^2 + \partial y^2}$). Do not worry if this is not clear to you at this point, it will become much clearer as you go through the course and study directional derivatives and the Sobel filter. 

Run the next cell to see the Sobel filter applied to the image hands. Change the argument axis to see the effect.

```sos
%use sos

#Apply filter
hands_edge = ndi.sobel(hands, axis= 0, mode='reflect')
# Compare
plt.close('all')
ndi_sobel_viewer = viewer([hands, hands_edge], title = ['Original', 'Edges'])
```

### Rotation in Scikit Image

The basic syntax of the function for grayscale images is ([see full doc](https://scikit-image.org/docs/stable/api/skimage.transform.html?highlight=rotate#skimage.transform.rotate)):
```python
output = skimage.transform.rotate(image, angle, resize=False, center=None, mode='reflect', preserve_range = True)
```

The parameters are:
* `input` (NumPy array): The source image, 
* `angle` (float): Rotation angle in degrees,
* `resize` (boolean): Specify whether to resize the image to contain the full original. Defaults to `False`
* `center` (iterable of length 2): Specifies the center of rotation. The order is (cols, rows).
* `mode` (string): Specify boundary conditions. We will use `'reflect'` throughout the course the course to mean mirroring boundary conditions when values outside the image are needed. Other options include `constant`, `nearest`, or `wrap'`.

In Scikit Image, we have a new parameter that nor OpenCV nor Scipy NDImage have. Scikit Image lets you specify the center of rotation, which for some applications can come as a very handy tool.

We will apply Scikit Image's rotation to `car` again, this time for an angle of $60^\circ$. Moreover, we will change the center of rotation. Run the next cell to see the result. Modify the cell and experiment with the different parameters.

```sos
%use sos

# First, we import the module transform
from skimage import transform

# Now we apply operation
rot_car_skimage = skimage.transform.rotate(car, angle = 60, resize=False, center=(50, 50), mode='reflect', preserve_range = True)

# And visualize
plt.close('all')
skimage_rotate_viewer = viewer([car, rot_car_skimage], title = ['Original', 'Rotated'], widgets = True)
```

### Edge Detection in Scikit Image

The basic syntax of the Sobel filter for grayscale images is ([see full doc](https://scikit-image.org/docs/dev/api/skimage.filters.html#skimage.filters.sobel)):

```python
output = skimage.filters.sobel(image, axis=None, mode='reflect')
```

The parameters are:
* `image` (NumPy array): The source image, 
* `axis` (int): If given, compute the axis in the specified direction. If `None`, the magnitude of the gradient is calculated,
* `mode` (string): Specify boundary conditions (see example above). 

As you can see, in this function, Scikit image gives more freedom to the user (you can produce both the magnitude of the gradient, and the gradient in each individual axis). Moreover, if you need a quicker and simpler way, Scikit-image has the filters `sobel_h` and `sobel_v`, that without any extra parameters apply the corresponding (horizontal or vertical) filter.

Run the next cell to apply Scikit-Image Sobel filter on the image `plate`. Modify the cell and play with the different values and with different images.

```sos
%use sos

# First, we import the module filters
from skimage import filters

# Apply filter
plate_edges_skimage = skimage.filters.sobel(plate, axis=None, mode='reflect')

# Visualize
plt.close('all')
skimage_edge_viewer = viewer([plate, plate_edges_skimage], title = ['Original', 'Edges'], widgets = True)
```

As you have realized by this brief introduction, no library is perfect, entirely complete, nor generally better than the others. For some applications and for some functions, you will find one library to be more useful than the others. In some cases, some libraries will not have the functions you need! So in complex applications, you may need to use two or even the three of them. In general, it is good to be aware of your options. 

###  <a name="exercise_iplibraries"></a> Try it yourself! (2 Points)

Congratulations! You are almost at the end of the introductory lab, this is your last task. It is your time to explore the three libraries (explore documentation, tutorials, etc.). A good access point is to go to their website (see the links in their names the first time we mentioned them). For **2 points**, then, choose one of the libraries and implement a **Gaussian smoothing (or blurring)** on the image `boats`. Modify the cell below, and store your result in the variable `blurred_boats`. Use:
* $\sigma = 5$,
* `'reflect'` (or `cv.BORDER_REFLECT`) boundary conditions, 
* and use default kernel size. 

If applicable, use: 
* `preserve_range = True`, and/or 
* `ksize = (0, 0)`.

<!-- grade_id: cell-3aad12f740231a32, points: 0, grade: false, solution: true -->
```sos
%use sos

blurred_boats = np.copy(boats)
# YOUR CODE HERE
```

<!-- grade_id: cell-1919408b3ba0b1c7, points: 2, grade: true, solution: false -->
```sos
%use sos
# Visualize 
plt.close('all')
gaussian_viewer = viewer([boats, blurred_boats], title = ['Original', 'Blurred']) 
```